Before you turn this project in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

# Load Data

The following code loads in the MNIST dataset and displays a few images and flattens the images and sets some autograder variables.

```python
from IPython.display import display
import numpy as np
from keras.preprocessing.image import array_to_img, img_to_array
from keras.datasets import mnist
from keras.utils.np_utils import to_categorical
import tensorflow as tf
sess = tf.InteractiveSession()
%load_ext autoreload
%autoreload 2

[X_mnist, y_mnist], _ = mnist.load_data()
for x in X_mnist[:5]:
    x = np.expand_dims(x, axis=-1)
    img = array_to_img(x)
    display(img)

X_mnist = X_mnist.reshape([60_000, 28*28]) / 255.
Y_mnist = to_categorical(y_mnist)
X_mnist, Y_mnist = X_mnist[:50], Y_mnist[:50]

M, N = X_mnist.shape
C = np.unique(y_mnist).shape[0]
H = 16

def passed(): print('✅')
```

# Task

- Define a one-hidden layer perceptron class called `TFMLP` with TensorFlow as in the computational graph

![](images/mlp_predict.svg)

<!-- grade_id: mlp_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Input Tests

<!-- grade_id: input_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import TFMLP

mlp = TFMLP(nb_feature=N, nb_hidden=H, nb_class=C)
init = tf.global_variables_initializer()
sess.run(init) # initialize tensorflow variables

assert mlp.X.op.type == 'Placeholder'
assert mlp.Y.op.type == 'Placeholder'
mlp.X.get_shape().as_list() == [None, 784]
mlp.Y.get_shape().as_list() == [None, 10]

passed()
```

# Variable Tests

<!-- grade_id: variable_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import TFMLP

mlp = TFMLP(nb_feature=N, nb_hidden=H, nb_class=C)
init = tf.global_variables_initializer()
sess.run(init) # initialize tensorflow variables

assert mlp.W1.op.type == 'VariableV2'
assert mlp.b1.op.type == 'VariableV2'
assert mlp.W1.get_shape().as_list() == [784, 16]
assert mlp.b1.get_shape().as_list() == [16]
assert mlp.W2.op.type == 'VariableV2'
assert mlp.b2.op.type == 'VariableV2'
assert mlp.W2.get_shape().as_list() == [16, 10]
assert mlp.b2.get_shape().as_list() == [10]

passed()
```

# Output Tests

<!-- grade_id: output_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import TFMLP

mlp = TFMLP(nb_feature=N, nb_hidden=H, nb_class=C)
init = tf.global_variables_initializer()
sess.run(init) # initialize tensorflow variables

assert mlp.loss.get_shape().as_list() == []
assert mlp.P.get_shape().as_list() == [None, 10]
assert mlp.accuracy.get_shape().as_list() == []

passed()
```

# Prediction Tests

<!-- grade_id: prediction_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import TFMLP

mlp = TFMLP(nb_feature=N, nb_hidden=H, nb_class=C)
init = tf.global_variables_initializer()
sess.run(init) # initialize tensorflow variables

y_pred = mlp.predict(X_mnist)
nb_train_ = len(X_mnist)
assert y_pred.shape == (M, C)

passed()
```

# Evaluation Test

<!-- grade_id: evaluation_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import TFMLP

mlp = TFMLP(nb_feature=N, nb_hidden=H, nb_class=C)
init = tf.global_variables_initializer()
sess.run(init) # initialize tensorflow variables

acc = mlp.evaluate(X_mnist, Y_mnist)
assert type(acc) == tf.float32
assert 0 <= acc <= 1

passed()
```

# Task

- Implement a `TFMLPWithGDOptimizer` class which performs optimization via gradient descent and extends your `TFMLP` class

<!-- grade_id: optimizer_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Gradient Descent Optimizer Test

<!-- grade_id: optimizer_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import TFMLPWithGDOptimizer, TFMLP

assert issubclass(TFMLPWithGDOptimizer, TFMLP)
mlp = TFMLPWithGDOptimizer(nb_feature=N, nb_hidden=H, nb_class=C)
init = tf.global_variables_initializer()
sess.run(init) # initialize tensorflow variables

X_sample, Y_sample = X_mnist[:50], Y_mnist[:50]
acc = mlp.evaluate(X_sample, Y_sample)
loss = mlp.forward(X_sample, Y_sample)
for _ in range(10):
    mlp.fit(X_sample, Y_sample, sess, nb_epoch=10)
    assert mlp.forward(X_sample, Y_sample) < loss
    loss = mlp.forward(X_sample, Y_sample)
    
assert mlp.evaluate(X_sample, Y_sample) > acc

passed()
```

# Ignore Cell Below

<!-- grade_id: ignore_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```